# Data Science - Unit 1 Sprint 2 Module 2

## Hypothesis Testing - Chi square tests

### Objectives

* Objective 01 - explain the purpose of a chi-square test and identify applications
* Objective 02 - set up a chi-square test for independence on two categorical variables
* Objective 03 - use a chi-square test p-value to draw the correct conclusion about the null and alternative hypothesis
* Objective 04 - explain how to compare the null hypothesis and sample and how randomness in the sampling process affect the result

#### Total notebook points: 9

### Introduction

Slasher movies thrill audiences by portraying a lone antagonist (typically male) who attacks innocent victims with extreme violence and without apparent motive. However, this exciting (if gory) sub-genre of horror film is criticized by those who view the violence as being used to “punish” female characters who engage in sexual activity during the film. 

### Dataset description:

To test this claim, study authors randomly sampled 50 North American slasher films released from 1960 to 2009 and coded the 485 characters appearing in them as being male or female, involved in sexual activity or not, and if they survived the film or not. 

The data appears in the `Slasher.csv` dataset. In this project, we are going to answer a slightly simpler question: Does a character's gender impact his or her chances of surviving a slasher film, or are the two factors unrelated?

You can find the `Slasher.csv` file and more information [here](https://github.com/LambdaSchool/data-science-practice-datasets/tree/main/unit_1/Slasher)

*Source: A. Welsh (2010). "On the Perils of Living Dangerously in the Slasher Horror Film: Gender Differences in the Association Between Sexual Activity and Survival," Sex Roles, Vol. 62 pp. 762-773*

**Task 1** - Load the data

Let's load the data! The URL has been provided as well as the imports for pandas and numpy.

* Load your CSV file into a DataFrame named `df_slasher`

<!-- grade_id: cell-fd8de6314d149c03, points: 0, grade: false, solution: true -->
```python
# Task 1

# Imports 
import pandas as pd
import numpy as np

data_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Slasher/Slasher.csv'

# YOUR CODE HERE
raise NotImplementedError()

# Print out your DataFrame
df_slasher.head()
```

```python
df_slasher
```

**Task 1 Test**

<!-- grade_id: cell-6526188acaf48e23, points: 1, grade: true, solution: false -->
```python
# Task 1 - Test

assert isinstance(df_slasher, pd.DataFrame), 'Have you created a DataFrame named `df_slasher`?'
assert len(df_slasher) == 485

```

**Task 2**- Character percentages

In this dataset, the variable codes are:

* ID = Character ID
* Gender = Character gender.  0 = male, 1 = female
* Activity = Character was portrayed as engaged in sexual activity.  0 = no, 1 = yes
* Survival = Character survival in the film. 0 = did not survive, 1 = survived

Your task is to calculate the frequency and relative frequency of male and female characters in this set of slasher films.

* Compute the frequencies for "Gender" for this dataset; call your table `gender_frequencies`
* Compute the relative frequency of female characters and call this value `f_percent`
* Compute the relative frequency of male characters and call this value `m_percent`

<!-- grade_id: cell-0d34d020c7d319e6, points: 0, grade: false, solution: true -->
```python
# Task 2

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 2 Test**

<!-- grade_id: cell-999f643c0b047fe5, points: 1, grade: true, solution: false -->
```python
# Task 2 - Test

assert isinstance(gender_frequencies, pd.Series), 'Have you calculated the gender frequencies?'

```

**Task 3** - Character survival

Calculate the frequency and relative frequency of character survival in this set of slasher films.

* Compute the frequencies for "Survival" for this dataset; call your table `survival_frequencies`
* Compute the frequency of character survival and call this value `survival`. Make sure to multiply by 100 so your value is a percent.

<!-- grade_id: cell-377c90fe789769b9, points: 0, grade: false, solution: true -->
```python
# Task 3

# YOUR CODE HERE
raise NotImplementedError()

# Print the percentage of characters that survive
print('Percentage of characters that survive:', survival)
```

**Task 3 Test**

<!-- grade_id: cell-6d42ea79de460476, points: 1, grade: true, solution: false -->
```python
# Task 3 - Test

assert isinstance(survival_frequencies, pd.Series), 'Have you created the frequency table?'

```

**Task 4**- Joint distribution: character survival and character sex

Next, we're going to calculate the joint distribution of *character sex ("Gender")* and movie *survival*. 

* Compute the joint distribution for character sex and survival and name your DataFrame `survival_fm_crosstab`. 
    * **Make sure to use the following order of your variables: "Gender", "Survival"**

*Hint: If you are working from the in-class examples, it will be helpful to include character sex as the index variable. But play around with the arguments of `pd.crosstab()` to get a better idea of your table.*

* Added Challenge: Rename your columns and index for clarity's sake.

>Gender: 0 = male, 1 = female

>Survival: 0 = did not survive, 1 = survived

<!-- grade_id: cell-71d255941aab5bf3, points: 0, grade: false, solution: true -->
```python
# Task 4

# YOUR CODE HERE
raise NotImplementedError()

# Print the table
survival_fm_crosstab
```

**Task 4 Test**

<!-- grade_id: cell-aab85db204edeb91, points: 1, grade: true, solution: false -->
```python
# Task 4 - Test

assert isinstance(survival_fm_crosstab, pd.DataFrame), 'Have you created the crosstab DataFrame?'

```

**Task 5** - Joint distribution with margins

Now we're going to find the joint distribution but this time add in the **table margins**.

* Compute the joint distribution and assign this to a variable called `survival_fm_margins`.

<!-- grade_id: cell-d71d9eea51fbf62e, points: 0, grade: false, solution: true -->
```python
# Task 5

# YOUR CODE HERE
raise NotImplementedError()

# Print the table
print(survival_fm_margins)
```

**Task 5 Test**

<!-- grade_id: cell-a3c7dedeca08a0cb, points: 1, grade: true, solution: false -->
```python
# Task 5 Test

assert isinstance(survival_fm_margins, pd.DataFrame), 'Did you create your joint distribution DataFrame?'
assert survival_fm_margins.shape == (3,3), "Did you include the margins in your distribution?"

# NO hidden tests for this task
```

**Task 6** - Joint distribution with normalization

Okay, now let's look at the distribution using the `normalize` argument. We want to compare the *relative* rate of survival. Calculate the distribution of survival conditional on character sex.

* Compute the joint distribution for character sex and survival and name your DataFrame `survival_fm_crosstab_norm` **and** normalize on the "index"
* Multiply the resulting DataFrame by 100 to return answers as percents
* Assign the *survival* percentage of female characters to the variable `survival_f`

<!-- grade_id: cell-94b846e0f979a986, points: 0, grade: false, solution: true -->
```python
# Task 6

# YOUR CODE HERE
raise NotImplementedError()

# Print the percentage of female characters that survive
print(survival_fm_crosstab_norm)
print('Female character survival:', survival_f)
```

**Task 6 Test**

<!-- grade_id: cell-1927d60ae3b59771, points: 1, grade: true, solution: false -->
```python
# Task 6 - Test

assert isinstance(survival_fm_crosstab_norm, pd.DataFrame), 'Have you created the crosstab DataFrame?'
assert survival_fm_crosstab_norm.all().min(axis=0) >= 1.0, "Did you multiply your values by 100?"

```

**Task 7 -** Statistical hypotheses

From the list of choices below, select the null and alternative hypotheses to test for a statistically significant relationship between character sex and survival.  Specify your answer in the next code block using `Answer = `.  For example, if the correct answer is choice B, you'll type `Answer = 'B'`.

A: $H_0$: Character sex and survival are related vs. $H_a$: Character sex and survival are unrelated$

B: $H_0$: Character sex and survival are unrelated vs. $H_a$: Character sex and survival are related$

C: $H_0: \mu_1 \neq \mu_2$ vs. $H_a: \mu_1 = \mu_2$

D: $H_0: \mu_1 = \mu_2$ vs. $H_a: \mu_1 \neq \mu_2$

<!-- grade_id: cell-da6613ef2df1b212, points: 0, grade: false, solution: true -->
```python
# Task 7

# YOUR CODE HERE
raise NotImplementedError()

```

**Task 7 Test**

<!-- grade_id: cell-37e2b8a5594319fc, points: 1, grade: true, solution: false -->
```python
# Task 7 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 8** - Chi-square test

Perform the chi-square test above using `stats.chi2_contingency()`. The docs for this function are found [here](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.chi2_contingency.html).

* Compute the chi-square test and assign the p-value to the variable `p_val`
* Assign the other variables returned by the function to any variable names you prefer


<!-- grade_id: cell-08a7fbc7c34ec943, points: 0, grade: false, solution: true -->
```python
# Task 8

# Import the scipy.stats function
from scipy.stats import chi2_contingency

# YOUR CODE HERE
raise NotImplementedError()

# Print out the p-value
print('The p-value is:', p_val)
```

**Task 8 Test**

<!-- grade_id: cell-f09cc2e930c43388, points: 1, grade: true, solution: false -->
```python
# Task 8 Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 9**

Select the correct conclusion at the 0.05 significance level from the list of choices below. Specify your answer in the next code block using `Answer = `.  For example, if the correct answer is choice B, you'll type `Answer = 'B'`.

A: We reject the null hypothesis at the 0.05 significance level and conclude that character survival and sex are related.

B: We fail to reject the null hypothesis at the 0.05 significance level and conclude that character survival and sex are related.

C: We reject the null hypothesis at the 0.05 significance level and conclude that character survival and sex are unrelated.

D: We fail to reject the null hypothesis at the 0.05 significance level and conclude that character survival and sex are unrelated.


<!-- grade_id: cell-52e04ab5206eb674, points: 0, grade: false, solution: true -->
```python
# Task 9

# YOUR CODE HERE
raise NotImplementedError()

```

**Task 9 Test**

<!-- grade_id: cell-9e8a26c1f374c2fa, points: 1, grade: true, solution: false -->
```python
# Task 9 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 10** - Create a bar plot

Create a side-by-side bar plot to illustrate the relationship between character sex and survival in slasher movies.  Sum up your overall results in a sentence or two.

This task will not be autograded - but it is part of completing the project.

<!-- grade_id: cell-a29031ed910c3c6c, points: 0, grade: true, solution: true -->
```python
# Task 10

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fc22cf60de1ac470, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE