## Autograded Notebook (Canvas & CodeGrade)

This notebook will be automatically graded. It is designed to test your answers and award points for the correct answers. Following the instructions for each Task carefully.

### Instructions

* **Download this notebook** as you would any other ipynb file
* **Upload** to Google Colab or work locally (if you have that set-up)
* **Delete `raise NotImplementedError()`**
* Write your code in the `# YOUR CODE HERE` space
* **Execute** the Test cells that contain `assert` statements - these help you check your work (others contain hidden tests that will be checked when you submit through Canvas)
* **Save** your notebook when you are finished
* **Download** as a `ipynb` file (if working in Colab)
* **Upload** your complete notebook to Canvas (there will be additional instructions in Slack and/or Canvas)

# Lambda School Data Science - Unit 1 Sprint 2 Module 1

## Hypothesis Testing - One and two-sample t-Tests

### Objectives

* Explain the purpose of a t-test and identify applications
* Use a t-test for independence to test for a statistically significant association between two categorical variables
* Use a t-test p-value to draw the correct conclusion about the null and alternative hypothesis

#### Total notebook points: 9

### Introduction

Mosquito nets have traditionally been an important tool to prevent mosquito bites in parts of the world where malaria is endemic. However, it may not be practical for an army that is on the move to set up and carry mosquito nets each night and day. Impregnating soldiers’ uniforms with insect repellent solves the mobility problem but also has drawbacks. First, the insect repellent quickly becomes ineffective with repeated washing and ironing and must be frequently reapplied. Second, in hot and humid climates the insect repellent can be absorbed through the skin, and the long-term effects of this exposure are unknown. One compromise is to have soldiers apply patches treated with insect repellent to their clothing. These patches would last longer because they would not be washed or ironed and would not expose the entire body to the insect repellent.

### Dataset Description

The `Mosquito.xlsx` dataset contains data recorded in an experiment conducted on male soldiers in the Indian Army who were stationed in the Tezpur/Solmara garrison in Northeast India. Thirty soldiers were randomly selected to receive one of three types of mosquito single repellent patch. After giving informed consent, the study participants affixed the patches at predetermined points on their uniforms and research assistants (who were blinded to the type of repellent used) counted the number of times a mosquito landed on each individual in an hour. 

Medical officers with the Indian Army have recorded data on mosquito bites and related illness for many years and can say with authority that the mean number of mosquito touches for soldiers not wearing any mosquito repellent is 8.2 per hour.**We wish to determine if wearing a single repellent patch changes the mean number of mosquito touches for soldiers compared to not wearing any mosquito repellant.**

*Adapted from: A. Bhatnagar and V.K. Mehta (2007). "Efficacy of Deltamethrin and Cyfluthrin Impregnated Cloth Over Uniform Against Mosquito Bites," Medical Journal Armed Forces India, Vol. 63, pp. 120-122.*

**Task 1** - Load the data

Let's load the data! The URL has been provided as well as the imports for pandas and numpy.

* load your CSV file into a DataFrame named `df_mosquito`

<!-- grade_id: cell-dae3789e00dd9b29, points: 0, grade: false, solution: true -->
```python
# Task 1
import pandas as pd
import numpy as np

# URL for the dataset
data_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Mosquito/Mosquito.csv'

# YOUR CODE HERE
raise NotImplementedError()

# Print out your DataFrame
df_mosquito.head()
```

**Task 1 Test**

<!-- grade_id: cell-a347d8dd7572e6d4, points: 1, grade: true, solution: false -->
```python
# Task 1 - Test

assert isinstance(df_mosquito, pd.DataFrame), 'Have you created a DataFrame named `df_mosquito`?'
assert len(df_mosquito) == 90

```

**Task 2 -** Statistical hypotheses

Write the null and alternative hypotheses using the experiment information described above.

This task will not be autograded - but it is part of completing the challenge.

**Task 2 ANSWER:**

FILL IN YOUR ANSWERS BELOW

Null hypothesis: 

Alternative hypothesis: 

$H_0: \mu =$ 

$H_a: \mu \neq$ 

**Task 3** - Calculate the mean

* Calculate the mean number of mosquito touches in the sample. Assign your answer to the variable `mosquito_touch_mean`.

<!-- grade_id: cell-22ced5cd2cae6230, points: 0, grade: false, solution: true -->
```python
# Task 3

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 3 Test**

<!-- grade_id: cell-3970fbb1a3eaefb5, points: 1, grade: true, solution: false -->
```python
# Task 3 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 4** - Calculate the standard deviation

* Calculate the standard deviation of the number of mosquito touches in the sample. Assign your answer to `mosquito_touch_std`.

<!-- grade_id: cell-9c4fe58a9fe2d7aa, points: 0, grade: false, solution: true -->
```python
# Task 4

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 4 Test**

<!-- grade_id: cell-4939c4289c39b4e3, points: 1, grade: true, solution: false -->
```python
# Task 4 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 5** - One-sample t-test

* Conduct a 1-sample t-test to test your hypotheses. Assign your t-test result to the variable `mosquito_pval`.

*Hint: The `stats.ttest_1samp()` function returns two values; read the documentation to select the p-value*

<!-- grade_id: cell-44b3f2a39081c5df, points: 0, grade: false, solution: true -->
```python
# Task 5

# Use the 'ttest_1samp' from the stats package
from scipy import stats

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 5 Test**

<!-- grade_id: cell-3be2eb1f16e366e6, points: 1, grade: true, solution: false -->
```python
# Task 5 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 6**

Report your conclusion at the 0.05 significance level. Write your answer in the cell below.

This task will not be autograded - but it is part of the project!

**Task 6 ANSWER:**

SHORT ANSWER HERE

## Use the following information to complete Tasks 7-13

### Introduction

More than 14,000 people finished the 2020 Disney Marathon held on January 12. The results by age and gender group are included in the `Disney.csv dataset`. 

**We wish to determine if the mean finishing time for male and female marathon runners is the same or if there is a difference in the mean finishing time between male and female marathon runners.**


[Source: Track Shack. 2020 Disney Marathon Race Results](https://www.trackshackresults.com/disneysports/results/wdw/wdw20/mar_results.php)

**Task 7** - Load the next dataset

Let's load the data! The URL has been provided.

* load your CSV file into a DataFrame named `df_disney`

<!-- grade_id: cell-2b389fa658a51cd2, points: 0, grade: false, solution: true -->
```python
# Task 7

# URL for Disney marathon dataset
data_url2 = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Disney_Marathon/Disney.csv'

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 7 Test**

<!-- grade_id: cell-c1f16b75c543e0d1, points: 1, grade: true, solution: false -->
```python
# Task 7 - Test

assert isinstance(df_disney, pd.DataFrame), 'Have you created a DataFrame named `df_disney`?'
assert len(df_disney) == 14106

```

**Task 8** - Write the hypothesis

Write the null and alternative hypotheses for the scenario described above in words and symbols.

This task will not be autograded - but it is part of completing the challenge.

**Task 8 ANSWER:**

FILL IN YOUR ANSWERS BELOW

Null hypothesis: 

Alternative hypothesis: 

$H_0: \mu =$ 

$H_a: \mu \neq$ 

**Task 9** - Create new series from a DataFrame

Create two separate Series (a pandas DataFrame column is a Series):

* one containing finishing times for male participants (`male_finish`)
* one containing finishing times for female participants (`female_finish`)

<!-- grade_id: cell-4290d3dd2f668a25, points: 0, grade: false, solution: true -->
```python
# Task 9

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 9 Test**

<!-- grade_id: cell-5384320ccd7f5048, points: 1, grade: true, solution: false -->
```python
# Task 9 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 10** - Calculate the mean finishing times

* Calculate the mean finishing time for male and female participants separately. Name your variables `male_finish_mean` and `female_finish_mean`.

<!-- grade_id: cell-f6c55ea4a2dba463, points: 0, grade: false, solution: true -->
```python
# Task 10

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 10 Test**

<!-- grade_id: cell-71bc6659db96de45, points: 1, grade: true, solution: false -->
```python
# Task 10 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 11** - Calculate the standard deviation

* Calculate standard deviation of the mean finishing time for male and female participants separately. Name your variables `male_finish_std` and `female_finish_std`.

<!-- grade_id: cell-3ac04d8339ec03c3, points: 0, grade: false, solution: true -->
```python
# Task 11

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 11 Test**

<!-- grade_id: cell-a2de7fe5b67a1886, points: 1, grade: true, solution: false -->
```python
# Task 11 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 12** - Conduct a 2-sample t-test

Conduct a 2-sample t-test to test your hypotheses:

* Assign the t-statistic to a variable called `disney_tval`
* Assign the p-value to a variable called `disney_pval`

Hint: The function returns two values and you can assign them with one line (example):

`variable1, variable2` = `some.function(arguments)`

<!-- grade_id: cell-d391e7e149994d87, points: 0, grade: false, solution: true -->
```python
# Task 12

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 12 Test**

<!-- grade_id: cell-35a1e3947ac6388a, points: 1, grade: true, solution: false -->
```python
# Task 12 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 13** - Report your result

Report your conclusion at the 0.05 significance level.

This task will not be autograded - but it is part of completing the challenge.

SHORT ANSWER HERE

### In your own words, answer the following questions:

These tasks will not be autograded - but they are part of completing the challenge.

**Task 14:** Explain the Central Limit Theorem.

SHORT ANSWER HERE

**Task 15:** Describe the Normal Distribution.

SHORT ANSWER HERE

**Task 16:** Describe the relationship between the Normal distribution and the t-distribution.

SHORT ANSWER HERE

**Task 17:** Write about who William Sealy Gosset was.

SHORT ANSWER HERE
